---
title: 'Intro to Pandas'
description: 'We have seen some of the fundamental tools needed in the python language, but we can build upon those to dive into pandas.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This exercise will reinforce the fundamenals of analyzing data using pandas.

`@instructions`
Assign two variables, a and b, and set the values such that the sum of a + b is 7.  

Save the result of sum of a + b in a new variable, c.

`@hint`
You an assign a variable with an `=`. For example, `x = 1`.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# assign the variables a and b

# assign the sum of a and b to a new variable, c
```

`@solution`
```{python}
a = 5
b = 2
c = a + b
```

`@sct`
```{python}

```
